### Codio Activity 22.4: Using `keras`

**Expected Time = 60 minutes**

**Total Points = 50**

This activity focuses on using the `keras` library to build an Artificial Neural Network for the titanic dataset.  You will now use all the rows of the data and the `age` and `fare` columns to build a basic network.  After building the model you will visualize the loss by epoch.

#### Index

- [Problem 1](#-Problem-1)
- [Problem 2](#-Problem-2)
- [Problem 3](#-Problem-3)
- [Problem 4](#-Problem-4)
- [Problem 5](#-Problem-5)

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
import tensorflow as tf

from tensorflow.keras.layers import Dense
from tensorflow.keras.models import Sequential
from tensorflow.keras.utils import plot_model

import warnings
warnings.filterwarnings('ignore')
```

Output:
```
2024-03-01 04:47:25.583303: E external/local_xla/xla/stream_executor/cuda/cuda_dnn.cc:9261] Unable to register cuDNN factory: Attempting to register factory for plugin cuDNN when one has already been registered
2024-03-01 04:47:25.583325: E external/local_xla/xla/stream_executor/cuda/cuda_fft.cc:607] Unable to register cuFFT factory: Attempting to register factory for plugin cuFFT when one has already been registered
2024-03-01 04:47:25.583923: E external/local_xla/xla/stream_executor/cuda/cuda_blas.cc:1515] Unable to register cuBLAS factory: Attempting to register factory for plugin cuBLAS when one has already been registered
2024-03-01 04:47:25.587640: I tensorflow/core/platform/cpu_feature_guard.cc:182] This TensorFlow binary is optimized to use available CPU instructions in performance-critical operations.
To enable the following instructions: AVX2 AVX_VNNI FMA, in other operations, rebuild TensorFlow with the appropriate compiler flags.
2024-03-01 04:47:26.035281: W tensorflow/com
```

```python
titanic = sns.load_dataset('titanic').dropna(subset = ['age'])
X = titanic[['age', 'fare']].values
y = titanic['survived'].values
print(X.shape)
```

Output:
```
(714, 2)
```

[Back to top](#-Index)

### Problem 1

#### A Basic Network

**10 Points**

To begin, use `keras` and the `Sequential` model to create a model with the following architecture:

- 1 Hidden Layer
- 1 Node in Hidden Layer
- 1 output node
- Sigmoid activation function on all nodes

This would be equivalent to the architecture in your first assignment.

<!-- grade_id: cell-2fad3e5bd06726d4, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)
single_node_model = ''
    
### BEGIN SOLUTION
tf.random.set_seed(42)
single_node_model = Sequential([
    Dense(1, activation = 'sigmoid'),
    Dense(1, activation = 'sigmoid')
    ])
### END SOLUTION

### ANSWER CHECK
single_node_model
```

Output:
```
2024-03-01 04:47:26.412588: I external/local_xla/xla/stream_executor/cuda/cuda_executor.cc:901] successful NUMA node read from SysFS had negative value (-1), but there must be at least one NUMA node, so returning NUMA node zero. See more at https://github.com/torvalds/linux/blob/v6.0/Documentation/ABI/testing/sysfs-bus-pci#L344-L355
2024-03-01 04:47:26.428098: I external/local_xla/xla/stream_executor/cuda/cuda_executor.cc:901] successful NUMA node read from SysFS had negative value (-1), but there must be at least one NUMA node, so returning NUMA node zero. See more at https://github.com/torvalds/linux/blob/v6.0/Documentation/ABI/testing/sysfs-bus-pci#L344-L355
2024-03-01 04:47:26.428240: I external/local_xla/xla/stream_executor/cuda/cuda_executor.cc:901] successful NUMA node read from SysFS had negative value (-1), but there must be at least one NUMA node, so returning NUMA node zero. See more at https://github.com/torvalds/linux/blob/v6.0/Documentation/ABI/testing/sysfs-bus-pci#L344-
```

```python
single_node_model.layers[0].activation
```

Output:
```
<function keras.src.activations.sigmoid(x)>
```

<!-- grade_id: cell-caad29a86aec42f8, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
single_node_model_ = Sequential([
    Dense(1, activation = 'sigmoid'),
    Dense(1, activation = 'sigmoid')
    ])
#
#
#
assert type(single_node_model.layers[0].activation) == type(single_node_model_.layers[0].activation)
assert type(single_node_model.layers[1].activation) == type(single_node_model_.layers[1].activation)

### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 2

#### Compiling the Network

**10 Points**

Now, use your `single_node_model` and specify the following elements using the `.compile` method.

- `optimizer = rmsprop`
- `loss = binary_crossentropy` or `bce`
- `metrics = ['accuracy']`

<!-- grade_id: cell-e253d21e36668ef0, points: 0, grade: false, solution: true -->
```python
### GRADED

tf.random.set_seed(42)
single_node_model.compile()


### BEGIN SOLUTION
tf.random.set_seed(42)
single_node_model.compile(optimizer='rmsprop',
                         loss = 'binary_crossentropy',
                         metrics = ['accuracy'])
### END SOLUTION

### ANSWER CHECK
single_node_model
```

Output:
```
<keras.src.engine.sequential.Sequential at 0x7ff2f6b06bf0>
```

<!-- grade_id: cell-f655e65a2446657f, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
single_node_model_.compile(optimizer='rmsprop',
                         loss = 'binary_crossentropy',
                         metrics = ['accuracy'])
#
#
#
assert single_node_model.loss == 'bce' or single_node_model.loss == 'binary_crossentropy'
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 3

#### Fit the model

**10 Points**

With your model setup, fit the model below using the following parameters:

- `epochs = 20`
- `batch_size = 10`
- `verbose = 0`

Also, leave the `tf.random.set_seed(42)` to assure proper grading.  Assign your fit model to the variable `history` below.

<!-- grade_id: cell-edc2684ae7362db5, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)
history = ''

    
### BEGIN SOLUTION
tf.random.set_seed(42)
history = single_node_model.fit(x = X, y = y, 
                                epochs = 20, 
                                batch_size=10,
                               verbose = 0)
### END SOLUTION

### ANSWER CHECK
history.params
```

Output:
```
2024-03-01 04:47:27.079070: I external/local_xla/xla/service/service.cc:168] XLA service 0x7ff1ac8b35f0 initialized for platform CUDA (this does not guarantee that XLA will be used). Devices:
2024-03-01 04:47:27.079089: I external/local_xla/xla/service/service.cc:176]   StreamExecutor device (0): NVIDIA GeForce RTX 3080 Ti Laptop GPU, Compute Capability 8.6
2024-03-01 04:47:27.082503: I tensorflow/compiler/mlir/tensorflow/utils/dump_mlir_util.cc:269] disabling MLIR crash reproducer, set env var `MLIR_CRASH_REPRODUCER_DIRECTORY` to enable.
2024-03-01 04:47:27.091095: I external/local_xla/xla/stream_executor/cuda/cuda_dnn.cc:454] Loaded cuDNN version 8904
WARNING: All log messages before absl::InitializeLog() is called are written to STDERR
I0000 00:00:1709268447.117241   23607 device_compiler.h:186] Compiled cluster using XLA!  This line is logged at most once for the lifetime of the process.
{'verbose': 0, 'epochs': 20, 'steps': 72}
```

<!-- grade_id: cell-f71d1269cb4ff9dc, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
history_ = single_node_model_.fit(x = X, y = y, 
                                epochs = 20, 
                                batch_size=10,
                               verbose = 0)

#
#
#
assert history.params == history_.params
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 4

#### Evaluate the model

**10 Points**

Use the `.evaluate` method of your `single_node_model` with the `X` and `y` arrays to examine the loss and accuracy of the model.  Assign these values as `single_loss` and `single_acc` below.

<!-- grade_id: cell-3fc8d4b03bcd60f1, points: 0, grade: false, solution: true -->
```python
### GRADED
single_loss = ''
single_acc = ''
    
### BEGIN SOLUTION
tf.random.set_seed(42)
results = single_node_model.evaluate(X, y)
single_loss = results[0]
single_acc = results[1]
### END SOLUTION

### ANSWER CHECK
print(single_loss)
print(single_acc)
```

Output:
```
23/23 [==============================] - 0s 2ms/step - loss: 0.6754 - accuracy: 0.5938
0.6753642559051514
0.593837559223175
```

<!-- grade_id: cell-0fb4c144e3ce33b4, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
results_ = single_node_model_.evaluate(X, y)
single_loss_ = results_[0]
single_acc_ = results_[1]

#
#
#

print(single_loss_)
print(single_acc_)
assert single_acc == single_acc_
### END HIDDEN TESTS
```

Output:
```
23/23 [==============================] - 0s 1ms/step - loss: 0.6755 - accuracy: 0.5938
0.6754515171051025
0.593837559223175
```

[Back to top](#-Index)

### Problem 5

#### A More Complex Model

**10 Points**

To try to improve the model, now build and evaluate a second model that uses a single hidden layer with 100 nodes, and a single output layer.  

For the hidden layer use the `relu` activation function and for the output layer use the `sigmoid` activation.  

Again, set the `np.random.seed(42)`, using the same compile settings and train settings for number of epochs and batch size.  Assign the accuracy of the model to `hundred_acc` below.

<!-- grade_id: cell-99ca08ba0f195e6e, points: 0, grade: false, solution: true -->
```python
### GRADED
complex_model = ''
hundred_acc = ''
    
### BEGIN SOLUTION
tf.random.set_seed(42)
complex_model = Sequential([
    Dense(100, activation = 'relu'),
    Dense(1, activation = 'sigmoid')
])
complex_model.compile(loss = 'bce', optimizer = 'rmsprop',
                     metrics = ['acc'])
np.random.seed(42)
history = complex_model.fit(X, y, epochs = 20, batch_size = 10,
                           verbose = 0)
hundred_acc = complex_model.evaluate(X, y)
### END SOLUTION

### ANSWER CHECK
print(hundred_acc)
```

Output:
```
23/23 [==============================] - 0s 957us/step - loss: 0.6348 - acc: 0.6765
[0.6347596645355225, 0.6764705777168274]
```

<!-- grade_id: cell-6abd29d1900b7ea3, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
complex_model_ = Sequential([
    Dense(100, activation = 'relu'),
    Dense(1, activation = 'sigmoid')
])
complex_model_.compile(loss = 'bce', optimizer = 'rmsprop',
                     metrics = ['acc'])
np.random.seed(42)
history_ = complex_model_.fit(X, y, epochs = 20, batch_size = 10,
                           verbose = 0)
hundred_acc_ = complex_model.evaluate(X, y)

#
#
#
assert hundred_acc == hundred_acc_
### END HIDDEN TESTS
```

Output:
```
23/23 [==============================] - 0s 930us/step - loss: 0.6348 - acc: 0.6765
```